# <span style="color:teal;">CIS 211 Project 5:  Building GUIs with Tk</span>

##### Due 11:00 P.M. Friday May 5

Reading:
* [TkDocs - Tk Tutorial](http://www.tkdocs.com/tutorial/index.html)
* [Tkinter Tutorial](http://www.python-course.eu/python_tkinter.php) from `python-course.eu`

### <span style="color:teal">Setup</span>

Execute the following code cell to tell IPython you want to create windows with Tk and to import the `tkinter` library.

```python
%gui tk
import tkinter as tk
```

<span style="color:red">**Important:**</span> &nbsp; The autograder tests assume `tkinter` has been imported as shown above.  You should follow the same convention, and refer to widget classes as `tk.Label`, `tk.Entry`, and so on.

<span style="color:red">**Important:**</span> &nbsp; The code cells where you define your GUIs contain assignment statements that create global variables.  For example, the first GUI is assigned to a variable named `hello_app`.  **Do not change these assignment statements**.  The autograder test cells will use these variables to access parts of your GUIs.

### <span style="color:teal">1. &nbsp; Hello, World (10 points)</span>

The code cell for this project has an assignment statement that creates a top level Tk application named `hello_app`.  

Add statements to the code cell so the program has a label that says "Hello, World":

<img src="http://www.cs.uoregon.edu/Classes/15S/cis211/images/hello_gui.png"/>

The example above uses 24-point Helvetica.  You can use the default text, but feel free to experiment with other fonts, font sizes, colors, *etc*.

The goal for this part of the project is just to make sure you have all the necessary software installed and can create a simple application.  There is no documentation requirement for this part.

##### <span style="color:red">Code:</span>

<!-- grade_id: hello, points: 4, grade: true, solution: true -->
hello_app = tk.Tk()

flabel = tk.Label(hello_app, text="Hello, World!", font = "Helvetica 24")
flabel.grid(row=0, column=0, sticky=tk.W, padx = 100, pady = 50)
pass

##### <span style="color:red">Autograder Tests:</span>

<!-- grade_id: hello_test_1, points: 2, grade: true, solution: false -->
```python
# The GUI should have a single widget
widgets = list(hello_app.children.values())
assert len(widgets) == 1
```

<!-- grade_id: hello_test_2, points: 2, grade: true, solution: false -->
```python
# Check to see if the widget is a Label
widgets = list(hello_app.children.values())
assert isinstance(widgets[0], tk.Label)
```

<!-- grade_id: hello_test_3, points: 2, grade: true, solution: false -->
```python
# The text in the label should start with "hello" (upper or lower case)
widgets = list(hello_app.children.values())
assert widgets[0]['text'].lower().startswith('hello')
```

Documentation:

    The point of part 1 is to get fimiliar with using Tkinter so we had to make a widget that just reads out "Hello, World!"

### <span style="color:teal">2. &nbsp; Payment Calculator (50 points)</span>

Create a program that computes the monthly payment on a loan. The payment is a function of the loan amount, the annual interest rate, and the number of years to pay off the loan. To calculate the payment, first compute two values named $r$ and $p$, defined by these equations:

$ r = ( \textrm{rate} \; / \; 100 ) \; / \; 12 $

$ p = 12 \times \textrm{years} $

The monthly payment is then

$$ \textrm{payment} = \frac{ r \times \textrm{amount} }{ 1 - (1 + r)^{-p} } $$

For example, if you take out a 30-year loan for \$150,000 at an annual rate of 4.5% your payment would be \$760.03 per month.

The GUI for this application should have four text entry boxes, and each box should have a label to its left that explains what the box holds.  The first three boxes are for the user to enter the loan amount, the interest rate, and the number of years to repay the loan. 

The GUI should have a button labeled “compute payment”. When the user clicks this button, your program should read the values of the three parameters, use the equations shown above to calculate the monthly payment, and display the payment amount in the fourth text entry box.

<img src="http://www.cs.uoregon.edu/Classes/15S/cis211/images/payment_gui.png"/>

The code cell for this program has the outline for the program.  First fill in the body of the function named `monthly_payment` and make sure it computes the correct payment value given an amount, rate, and loan period.  When your write the button callback function have it call the `monthly_payment` function.

Next add statements that create the labels, boxes, and button for the application.  When you create the Entry and Button widgets, make sure you save them in global variables with the following names so the autograder tests will succeed:
* `amount`
* `rate`
* `years`
* `payment`
* `button`

##### <span style="color:red">Code:</span>

<!-- grade_id: payment, points: 15, grade: true, solution: true -->
```python
def monthly_payment(amt, pct, yrs):
    '''
        (num,num,num) --> num
        
        This function computes the monthly payment of a loan given the input amount of a loan, percentage, and years.
        
        Returns payment, which is the monthly amount.
    '''
    # YOUR CODE HERE
    r= (pct/100)/12
    p= 12 * yrs
    payment = (r * amt)/(1-(1+r)**-p)
    return payment
    pass

def compute_payment_cb():
    '''
        ()--> None
        
        This function allows the Tkinter to use the function, monthly_payment
        
        No return statement so it returns nothing.
    '''
    a = float(amount.get())
    r = float(rate.get())
    y =  float(years.get())
    pay = monthly_payment(a,r,y)
    payment.delete(0, tk.END)
    payment.insert(0, "${:.2f}".format(pay))

payment_app = tk.Tk()

alabel = tk.Label(payment_app, text="Loan Amount:")
alabel.grid(row=1, column=1, sticky=tk.W, padx = 20, pady = 30)

rlabel = tk.Label(payment_app, text="Interest Rate:")
rlabel.grid(row=2, column=1, sticky=tk.W, padx = 20, pady = 30)

ylabel = tk.Label(payment_app, text="Loan Period:")
ylabel.grid(row=3, column=1, sticky=tk.W, padx = 20, pady = 20)

plabel = tk.Label(payment_app, text="Total Payment:")
plabel.grid(row=4, column=1, sticky=tk.W, padx = 20, pady = 20)
 
amount = tk.Entry(payment_app)
amount.grid(row=1, column=2, padx = 20)

rate = tk.Entry(payment_app)
rate.grid(row=2, column=2, padx = 20)

years = tk.Entry(payment_app)
years.grid(row=3, column=2, padx = 20)

payment = tk.Entry(payment_app)
payment.grid(row=4, column=2, padx = 20)

button = tk.Button(payment_app, text='Compute Payment', command=compute_payment_cb)
button.grid(row=5, column=2, columnspan=2, pady = 30)



```

##### <span style="color:red">Autograder Tests:</span>

<!-- grade_id: payment_test_1, points: 5, grade: true, solution: false -->
```python
# Test the callback function to make sure it is computing payments accurately
assert round(monthly_payment(20000, 3.0, 6), 2) == 303.87
```

<!-- grade_id: payment_test_2, points: 10, grade: true, solution: false -->
```python
# Verify the number of widgets of each type
counts = { tk.Label: 0, tk.Entry: 0, tk.Button: 0 }

for x in payment_app.children.values():
    counts[type(x)] += 1
    
assert counts == { tk.Label: 4, tk.Entry: 4, tk.Button: 1 }
```

<!-- grade_id: payment_test_3, points: 10, grade: true, solution: false -->
```python
# Store values in the GUI, click the button, check the result
amount.insert(0, '20000')
rate.insert(0, '3.0')
years.insert(0, '6')
button.invoke()

assert payment.get() == '$303.87'
```

##### <span style="color:red">Documentation:</span>

<!-- grade_id: doc_2, points: 10, grade: true, solution: true -->
This part of the project uses Tkinter and two functions, monthly_payment and compute_payment_cb. When someone runs the code, Tkinter pulls up a widget that allows you to put in the total amount of the loan with interest and how ever long it takes to pay off the loan in years. The function monthly payment does all of the calculations while compute_payment_cb basically connects monthly_payment and the widget from Tkinter together to allow to input everything and get a return total amount.

### <span style="color:teal">3. &nbsp; Canvas (40 points)</span>

Fill in the application named `canvas_app`.  The application will have just one widget, a Tk canvas with a size of $400 \times 400$ pixels. 

You also need to fill in the body of the function named `make_circle` that will draw a circle on your canvas.  The parameters are a reference to the canvas widget, three integers for the size and location of the circle, and a string with a color name.

Use the function to draw three circles on your canvas.  You can choose whatever size, location, and color you would like.

<img src="http://www.cs.uoregon.edu/Classes/16W/cis211/images/gui/three_circles.png"/>


##### <span style="color:red">Code:</span>

<!-- grade_id: canvas, points: 15, grade: true, solution: true -->
```python
def make_circle(canvas, radius, x, y, color):
    '''
        (canvas, num, num, ,num, str) --> canvas.create_oval
        
        This function is used to use Tkinter and brings up a canvas and allow Tk to create circles.
        
        returns canvas.create_oval
    '''
    return canvas.create_oval(x-rad,y-rad,x+rad,y+rad,width=0,fill=color)
    
canvas_app = tk.Tk()

drawing = tk.Canvas(canvas_app, height=400, width=400, background='white')
drawing.pack(padx = 10, pady = 10)
drawing.create_oval((250,250), (330,330), fill='green')
drawing.create_oval((150,150), (225,225), fill='yellow')
drawing.create_oval((25,25), (130,130), fill='green')


pass

```

##### <span style="color:red">Autograder Tests:</span>

<!-- grade_id: canvas_test_1, points: 5, grade: true, solution: false -->
```python
# See if the GUI has a Canvas widget
widgets = list(canvas_app.children.values())
assert len(widgets) == 1 and isinstance(widgets[0], tk.Canvas)
```

<!-- grade_id: canvas_test_2, points: 10, grade: true, solution: false -->
```python
# See if the canvas has circles on it
widgets = list(canvas_app.children.values())
for n in widgets[0].find_all():
    assert widgets[0].type(n) == 'oval'
```

##### <span style="color:red">Documentation:</span>

<!-- grade_id: doc_3, points: 10, grade: true, solution: true -->
This part uses Tkinter and the function make_circle to bring up canvas and display 3 different circles on the canvas. 